<a href="https://colab.research.google.com/github/atonendeley/CE880-Data-Science/blob/main/LabCourseWork/week7/CE880_Week_7_Course_Work1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Atonye Endeley"
COLLABORATORS = "2103195"
```

---

---
# Welcome to CE880
### This is your week-7 : Problem notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/). 

---
## Question 1: 

#### Bootstrapping

We are generating a random vector, which will be our population, whose law of distribution (expected to be the same as that of the population) is unknown. Write a function to pick a sample from our population and by applying the Bootstrapping procedure, where sample `size=300`: and `B=1000` number of re-sampling

<!-- grade_id: my_bootstrapping_ans, points: 0, grade: false, solution: true -->
```python
import numpy as np
import random
from sklearn.utils import resample
np.random.seed(123)
pop = np.random.randint(0,500 , size=100)
print(pop.mean()) # Print population mean
sample = np.random.choice(pop, size=300) #so n=300

def my_bootstrapping():

        """Write a function to pick a sample from our population and by applying the
    Bootstrapping procedure, where sample size=300: and B=10000 number of re-sampling """
    # YOUR CODE HERE

sample_mean = []
    

for _ in range(10000):  #so B=10000
  sample_n = np.random.choice(pop, size=300)
  sample_mean.append(sample_n.mean())


print(pop.mean())

    #raise NotImplementedError()
```

Output:
```
253.76
253.76
```

<!-- grade_id: my_bootstrapping_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_bootstrapping(), 253.241, rel_tol = 0.05)
```

Output:
```
TypeError                                 Traceback (most recent call last)
<ipython-input-2-609110f226a9> in <module>()
      1 # Check you solution by running this cell
      2 import math
----> 3 assert math.isclose(my_bootstrapping(), 253.241, rel_tol = 0.05)

TypeError: must be real number, not NoneType
```

---
## Question 2: 

#### Majority Voting
Given a list = `[1,2,3,1,2,1,1,1,3,2,2,1]`. Write a function to find the majority voting.

<!-- grade_id: my_pairedttest_ans, points: 0, grade: false, solution: true -->
```python
# Example of the Paired Student's t-test
a = np.array([1,2,3,1,2,1,1,1,3,2,2,1])

def my_majorityvote():
      """Write a function to find the majority voting."""
    # YOUR CODE HERE
count = {}
for i in a:
    if i not in count:
      count[i] = 1
    else:
        count[i] += 1
m_count = max(count.values())   

for key in count:
    if count[key] == m_count:
      print (key)    
    raise NotImplementedError()
```

Output:
```
1
NotImplementedError                       Traceback (most recent call last)
<ipython-input-18-5a30bae103c7> in <module>()
     16     if count[key] == m_count:
     17       print (key)
---> 18     raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: my_pairedttest_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_majorityvote(), 1, rel_tol = 0.05)
```

Output:
```
TypeError                                 Traceback (most recent call last)
<ipython-input-29-2dab1ce26478> in <module>()
      1 # Check you solution by running this cell
      2 import math
----> 3 assert math.isclose(my_majorityvote(), 1, rel_tol = 0.05)

TypeError: must be real number, not NoneType
```

---
## Question 3: 

#### Decision Tree
USe the IRIS dataset from `sklearn` to build a decision tree classifier. The dataset contains information for three classes of the IRIS plant, namely IRIS Setosa, IRIS Versicolour, and IRIS Virginica, with the following attributes: sepal length, sepal width, petal length, and petal width. Your aim is to train a decision tree on training data and return the test accuracy.

<!-- grade_id: my_DTree_ans, points: 0, grade: false, solution: true -->
```python
import pandas as pd
import numpy as np
from sklearn.datasets import load_iris
from sklearn.metrics import accuracy_score
from sklearn.tree import DecisionTreeClassifier
from sklearn import metrics
# Reading the Iris.csv file
data = load_iris()

# Extracting Attributes / Features
X = data.data

# Extracting Target / Class Labels
y = data.target

# Import Library for splitting data
from sklearn.model_selection import train_test_split

# Creating Train and Test datasets
X_train, X_test, y_train, y_test = train_test_split(X,y, random_state = 50, test_size = 0.25)

def my_DTree():
    """Write code to to train a decision tree classifier and return the test accuracy"""
        # YOUR CODE HERE
clf = DecisionTreeClassifier()

# Train Decision Tree Classifer

clf = clf.fit(X_train,y_train)

y_pred =  clf.predict(X_test)

 #Model Accuracy, how often is the classifier correct?
print("Accuracy:", metrics.accuracy_score(y_test, y_pred))

```

Output:
```
Accuracy: 0.9473684210526315
```

<!-- grade_id: my_DTree_test, points: 2, grade: true, solution: false -->
```python
# Check you solution by running this cell
import math
assert math.isclose(my_DTree(), 0.94736, rel_tol = 0.05)
```

Output:
```
TypeError                                 Traceback (most recent call last)
<ipython-input-9-f392ced02482> in <module>()
      1 # Check you solution by running this cell
      2 import math
----> 3 assert math.isclose(my_DTree(), 0.94736, rel_tol = 0.05)

TypeError: must be real number, not NoneType
```