<a href="https://colab.research.google.com/github/atonendeley/CE880-Data-Science/blob/main/LabCourseWork/week1/CE880_Week1_C_ourse_Work1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "ATONYE ENDELEY"
COLLABORATORS = "2103195"
```

---

---
# Welcome to CE880
### This is your week-1 : Problem 1 notebook

For this problem set, we'll be using the Jupyter notebook and please upload this notebook to [Google Colab](https://colab.research.google.com/). 


---
## Question 1:

Write a function that calculate the area of the square by using the formula as `area = side*side`
At last, print the area of the square to get the output.

<!-- grade_id: areaofSqr_ans, points: 0, grade: false, solution: true -->
```python
def areaofSqr(side):

    """Compute the area of squares with side = any number."""
    
    # YOUR CODE HERE
    area = side*side 
    return area
    side = 10
    print(areaofsqr(side))
    raise NotImplementedError()
```

```python
areaofSqr(10)
```

Output:
```
100
```

<!-- grade_id: areaofSqr_test, points: 1, grade: true, solution: false -->
```python
"""Check that squares returns the correct output for several inputs"""
assert areaofSqr(1) == 1
assert areaofSqr(2) == 4
assert areaofSqr(-3) == 9
assert areaofSqr(10) == 100
assert areaofSqr(11) == 121
```

---

## Question 2:

A `trapezoid` is a quadrilateral with two sides parallel. The `trapezoid` is equivalent to the British definition of the `trapezium`. An isosceles `trapezoid` is a `trapezoid` in which the base angles are equal so

<img src="https://www.w3resource.com/w3r_images/python-math-image-exercise-3.svg"
     alt="Markdown Monster icon"
     style="float: left; margin-right: 10px;" />

<!-- grade_id: areaofTrap_ans, points: 0, grade: false, solution: true -->
```python
def areaofTrap(base1, base2, height):
    """Compute the sum of the squares of numbers from 1 to n."""
    
    # YOUR CODE HERE
    # calculate the area
    area = 0.5 * (base1 + base2) * height
    print("areaofTrap = %.2f " %area)
    return area
    areaofTrap(9, 6, 4)

    
    raise NotImplementedError()
```

```python
areaofTrap(5, 3, -9)
```

Output:
```
areaofTrap = -36.00 
-36.0
```

<!-- grade_id: areaofTrap_test, points: 1, grade: true, solution: false -->
```python
"""Check that areaofTrap returns the correct answer for various inputs."""
assert areaofTrap(10, 10, 20) == 200.0
assert areaofTrap(5, 10, 20) == 150.0
assert areaofTrap(5, 3, 5) == 20.0
assert areaofTrap(5, 3, -9) == -36.0
```

Output:
```
areaofTrap = 200.00 
areaofTrap = 150.00 
areaofTrap = 20.00 
areaofTrap = -36.00
```

---
## Question 3: 

 Write a Python program to reverse a string. Example: Sample String: `1234xyz` Expected Output: `zyx4321`

<!-- grade_id: reversestring_ans, points: 0, grade: false, solution: true -->
```python
def reverse_string(str1):

  """Write your code here to reverse the string"""
     # YOUR CODE HERE

  string = ""   # Declaring empty string to store the reversed string  
  for i in str1:
    string = i + string
  return string
  str1 = "Happy"
  print("The reverse string is:", reverse_string(str1)) 

  raise NotImplementedError()
```

```python
reverse_string('CE880')
```

Output:
```
'088EC'
```

<!-- grade_id: reversestring_test, points: 1, grade: true, solution: false -->
```python
"""Check that areaofTrap returns the correct answer for various inputs."""
assert reverse_string('Colchester') == 'retsehcloC'
assert reverse_string('University') == 'ytisrevinU'
assert reverse_string('Data Science') == 'ecneicS ataD'
assert reverse_string('CE880') == '088EC'
```